## Module 4

#### In this assignment, you will continue working on the movie data from IMDB.
- The data includes movies and ratings from the IMDB website
- Data File(s): imdb.xlsx

#### Data file contains 3 sheets:
- “imdb”: contains records of movies and ratings scraped from IMDB website
- “countries”: contains the country (of origin) names
- “directors”: contains the director names

We have loaded and joined the data as "df" for you. Follow the instructions and finish the rest part.

<!-- grade_id: init_test, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import imp, os, sys
sol = imp.load_compiled("solutions", "./solutions.py")
sol.get_solutions("imdb.xlsx")
from nose.tools import assert_equal
from pandas.util.testing import assert_frame_equal, assert_series_equal
```

<!-- grade_id: init_user, points: 0, grade: true, solution: false -->
```python
# Loading the data
import pandas as pd

xls = pd.ExcelFile('imdb.xlsx')
df = xls.parse('imdb')
df_directors = xls.parse('directors')
df_countries = xls.parse('countries')

df = pd.merge(left=df, right=df_countries, 
              how='inner', left_on='country_id', 
              right_on='id')

df = pd.merge(left=df, right=df_directors, 
              how='inner', left_on='director_id', 
              right_on='id')

print("Finished.")
```

Output:
```
Finished.
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
""" Q1: 
Get the summary statistics for imdb_score and gross, then use the describe() function to summarize this visually. Save the
result in a variable called score_gross_description and print it.
"""

# your code here
df.describe()
score_gross_description = df[["imdb_score", "gross"]].describe()
score_gross_description
# type(score_gross_description)
```

Output:
```
imdb_score         gross
count  178.000000  1.780000e+02
mean     8.294382  1.030402e+08
std      0.266960  1.242549e+08
min      8.000000  8.060000e+03
25%      8.100000  1.318510e+07
50%      8.200000  5.194371e+07
75%      8.475000  1.522436e+08
max      9.300000  6.232795e+08
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
assert_frame_equal(score_gross_description, sol.score_gross_description)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
"""Q2:
What is the average rating of the director Christopher Nolan's movies? Save this value in a variable called nolan_mean and 
print.
"""

# your code here
nolan = df[df["director_name"] == "Christopher Nolan"]
nolan_mean = nolan["imdb_score"].mean()
nolan_mean
```

Output:
```
8.6
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
assert_equal(nolan_mean, sol.nolan_mean)
```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
"""Q3: 
Create a series called 'directors' that contains each director's name and his or her average rating.  Print out the type of your variable.
Use the 'directors' series to find the average rating for Steven Spielberg.  Print the value.
"""

# your code here
# df.head()
directors = df.groupby("director_name")["imdb_score"].mean()
type(directors)
directors

```

Output:
```
director_name
Akira Kurosawa           8.70
Alejandro Amenabar       8.10
Alejandro G. Inarritu    8.10
Alfred Hitchcock         8.50
Andrew Stanton           8.30
                         ... 
Victor Fleming           8.15
Vincent Paronnaud        8.00
Wes Anderson             8.10
Wolfgang Petersen        8.40
Woody Allen              8.10
Name: imdb_score, Length: 117, dtype: float64
```

```python
print(directors["Steven Spielberg"])
```

Output:
```
8.48
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
assert_series_equal(directors, sol.directors)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
"""Q4:
Select the non-USA movies made after 1960 by Hayao Miyazaki.
Save the result in a DataFrame called 'miyazaki', then print it.

Here are the steps:
1. Query the data ('df' DataFrame) based on the following conditions:
- Non-USA movies (country_id != 1)
- Movies made after 1960 (title_year > 1960)
- Movies made by director Hayao Miyazaki (director_id == 46)
2. Save the filtered data in a DataFrame called 'miyazaki' and print it

"""

# your code here
non_usa_movies = df["country_id"] != 1
# non_usa_movies
```

```python
after_1960_movies = df["title_year"] > 1960
# after_1960_movies
```

```python
hayao = df["director_id"] == 46
# hayao
```

```python
# miyazaki = df[(df["country_id"] != 1) & (df["title_year"] > 1960) & (df["director_id"] == 46)]
miyazaki = df[non_usa_movies & after_1960_movies & hayao]
miyazaki
```

Output:
```
movie_title  director_id  country_id content_rating  title_year  \
152         Spirited Away           46           4             PG        2001   
153     Princess Mononoke           46           4          PG-13        1997   
154  Howl's Moving Castle           46           4             PG        2004   

     imdb_score     gross  duration  id_x country  id_y   director_name  
152         8.6  10049886       125     4   Japan    46  Hayao Miyazaki  
153         8.4   2298191       134     4   Japan    46  Hayao Miyazaki  
154         8.2   4710455       119     4   Japan    46  Hayao Miyazaki
```

<!-- grade_id: q4_test, points: 2, grade: true, solution: false -->
```python
assert_frame_equal(miyazaki, sol.miyazaki)
print("Success!")
```

Output:
```
Success!
```

```python
df.head()
```

Output:
```
movie_title  director_id  country_id content_rating  \
0  The Shawshank Redemption           34           1              R   
1            The Green Mile           34           1              R   
2             The Godfather           33           1              R   
3    The Godfather: Part II           33           1              R   
4            Apocalypse Now           33           1              R   

   title_year  imdb_score      gross  duration  id_x country  id_y  \
0        1994         9.3   28341469       142     1     USA    34   
1        1999         8.5  136801374       189     1     USA    34   
2        1972         9.2  134821952       175     1     USA    33   
3        1974         9.0   57300000       220     1     USA    33   
4        1979         8.5   78800000       289     1     USA    33   

          director_name  
0        Frank Darabont  
1        Frank Darabont  
2  Francis Ford Coppola  
3  Francis Ford Coppola  
4  Francis Ford Coppola
```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
"""Q5: 
Create a Pivot Table that shows the median rating for each director, grouped by their respective countries. Name your variable
'pivot_agg'
"""

# your code here
pivot_agg = pd.pivot_table(df, index = ["country","director_name"], values = ["imdb_score"], aggfunc=["median"])
pivot_agg
```

Output:
```
median
                                  imdb_score
country      director_name                  
Argentina    Juan Jose Campanella       8.20
Australia    George Miller              8.10
Brazil       Fernando Meirelles         8.70
             Jose Padilha               8.10
Canada       Denis Villeneuve           8.20
...                                      ...
USA          Tony Scott                 8.00
             Victor Fleming             8.15
             Wes Anderson               8.10
             Woody Allen                8.10
West Germany Wolfgang Petersen          8.40

[125 rows x 1 columns]
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
assert_frame_equal(pivot_agg, sol.pivot_agg)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
"""Q6:
How long did the movie Gladiator aim to keep your attention? Save the series with this information
in a variable called 'gladiator_duration', then print it.
"""

# your code here
df.head()
gladiator = df[df["movie_title"] == "Gladiator"]
gladiator
```

Output:
```
movie_title  director_id  country_id content_rating  title_year  \
51   Gladiator           90           1              R        2000   

    imdb_score      gross  duration  id_x country  id_y director_name  
51         8.5  187670866       171     1     USA    90  Ridley Scott
```

```python
gladiator_duration = df[df["movie_title"] == "Gladiator"]["duration"]
# gladiator_duration = gladiator["duration"]
gladiator_duration
```

Output:
```
51    171
Name: duration, dtype: int64
```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
assert_series_equal(gladiator_duration, sol.gladiator_duration)
print("Success!")
```

Output:
```
Success!
```

```python

```